# Manipuler des tableaux

Dans cette feuille, vous allez apprendre à effectuer quelques
manipulations simples sur les tableaux, comme nous l'avions fait au
premier semestre avec les `vector` de C++. En Python, de tels tableaux
peuvent être représentés par les `array` de la bibliothèque `NumPy`
(usuellement abrégée en `np`):

```python
import numpy as np
```

## Tableaux à deux dimensions

Voilà un tableau à deux dimensions avec deux lignes et quatre
colonnes:

```python
T = np.array([[1, 2, 3, 4], 
              [5, 6, 7, 8]])
```

```python
T
```

Output:
```
array([[1, 2, 3, 4],
       [5, 6, 7, 8]])
```

On peut retrouver les tailles de ce tableau avec:

```python
T.shape
```

Output:
```
(2, 4)
```

Vous vous rappellez que les `vector` de C++ sont intrinsèquement des
tableaux à une dimension, et que l'on émule des tableaux à deux
dimensions avec des tableaux de tableaux. Ici, en revanche, les
tableaux `array` de numpy permettent de construire explicitement des
tableaux à deux dimensions.

### Exercice

1. Construire un tableau à trois lignes et trois colonnes, contenant
   les entiers de 1 à 9 de gauche à droite et de haut en bas comme
   dans la figure suivante:

       1 2 3
       4 5 6
       7 8 9

<!-- grade_id: cell-1b95ffffb22a38df, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
T2 = np.array([[1,2,3],[4,5,6],[7,8,9]])
T2
```

Output:
```
array([[1, 2, 3],
       [4, 5, 6],
       [7, 8, 9]])
```

Nous testons la forme du tableau:

<!-- grade_id: cell-3dcbc34e7cdb8449, points: 1, grade: true, solution: false -->
```python
assert T2.shape == (3,3)
```

ainsi que son contenu :

<!-- grade_id: cell-b2598336ebcb56a2, points: 1, grade: true, solution: false -->
```python
assert [ T2[i,j] for i in range(3) for j in range(3) ] == [1, 2, 3, 4, 5, 6, 7, 8, 9]
```

Voici comment accéder au contenu d'une case individuelle du tableau :

```python
T2[1,2]
```

Output:
```
6
```

Cette case est en deuxième ligne et troisième colonne: en effet, comme
en C++, les lignes et colonnes sont numérotées à partir de 0.

Si l'on veut extraire toute une ligne, ou toute une colonne, on
remplace la coordonnée que l'on ne veut pas spécifier par `:`.

Voici donc la deuxième colonne :

```python
T2[:,1]
```

Output:
```
array([2, 5, 8])
```

Extraire la deuxième ligne du tableau, et affectez-la à la variable
`l` dont vous afficherez le contenu :

<!-- grade_id: cell-0abb1396ab91dd95, points: 0, grade: false, solution: true -->
```python
l = T2[1,:]
l
```

Output:
```
array([4, 5, 6])
```

<!-- grade_id: cell-e22d3374bf68c443, points: 1, grade: true, solution: false -->
```python
assert isinstance(l, np.ndarray)
assert l.shape == (3,)
assert list(l) == [4,5,6]
```

## Tableaux à trois dimensions et plus

Pour le moment, nous avons utilisé des tableaux à deux dimensions.
Ultérieurement, notamment pour représenter des images, nous aurons
besoin de tableaux de plus grande dimension: un seul nombre ne suffit
en effet pas pour représenter un pixel.

`Numpy` permet de représenter des tableaux de toute dimension. Voici
un tableau de dimension 3 :

```python
T3D = np.array([[[ 1, 2, 3], [ 4, 5, 6], [ 7, 8, 9]],
                [[10,11,12], [13,14,15], [16,17,18]],
                [[19,20,21], [22,23,24], [25,26,27]]
                ])
```

On peut le voir comme un tableau à trois couches :

```python
T3D[:,:,0]
```

Output:
```
array([[ 1,  4,  7],
       [10, 13, 16],
       [19, 22, 25]])
```

```python
T3D[:,:,1]
```

Output:
```
array([[ 2,  5,  8],
       [11, 14, 17],
       [20, 23, 26]])
```

```python
T3D[:,:,2]
```

Output:
```
array([[ 3,  6,  9],
       [12, 15, 18],
       [21, 24, 27]])
```

### Exercices

Extraire la première colonne de la deuxième couche de `T3D` et
stockez-la dans la variable `C`:

<!-- grade_id: cell-81c4d71ee7b1a54d, points: 0, grade: false, solution: true -->
```python
C = T3D[:,0,1]
C
```

Output:
```
array([ 2, 11, 20])
```

Notez que c'est un tableau à une dimension, donc noté en ligne !

<!-- grade_id: cell-995993eb1b58947e, points: 1, grade: true, solution: false -->
```python
assert list(C) == [2, 11, 20]
```

Extraire un tableau contenant la première colonne de chacune des trois
couches de `T3D` et stockez le dans la variable `C`. Notez que l'on
souhaite que ces colonnes soient bien représentées par des colonnes
dans `C` !

<!-- grade_id: cell-abdcb2e9f73d6dfe, points: 0, grade: false, solution: true -->
```python
C = np.array(T3D[:,0,:])
C
```

Output:
```
array([[ 1,  2,  3],
       [10, 11, 12],
       [19, 20, 21]])
```

<!-- grade_id: cell-3227c3665bd8802d, points: 1, grade: true, solution: false -->
```python
for i in range(3):
    assert np.array_equal(T3D[:,0,i], C[:,i])
```

## Statistiques simples sur les tableaux

Numpy permet de faire des statistiques simples sur les
tableaux. Revenons à notre tableau `T` :

```python
T
```

Output:
```
array([[1, 2, 3, 4],
       [5, 6, 7, 8]])
```

Calculez à la main :
- la moyenne de chaque colonne de `T`;
- la moyenne de chaque ligne de `T`;
- la moyenne de tous les éléments du tableau `T`.

```python
def moyennecol(i : int,Tab):
    return sum(Tab[:,i])/len(Tab[:,i])
def moyennelig(i: int,Tab):
    return sum(Tab[i,:])/len(Tab[i,:])
def moy(Tab):
    return sum(sum(Tab))/(len(Tab)*len(Tab[0,:]))
moyscol = [moyennecol(i,T) for i in range(0,len(T[0,:]))]
moyslig = [moyennelig(i,T) for i in range(0,len(T))]
moyenne = moy(T)
moyscol,moyslig,moyenne
```

Output:
```
([3.0, 4.0, 5.0, 6.0], [2.5, 6.5], 4.5)
```

Comparez vos résultats avec ceux des calculs suivants :

```python
T.mean(axis=0)
```

Output:
```
array([3., 4., 5., 6.])
```

```python
T.mean(axis=1)
```

Output:
```
array([2.5, 6.5])
```

```python
T.mean()
```

Output:
```
4.5
```

## Conclusion

Voilà vous avez vu tous les éléments de manipulation des tableaux
`NumPy` dont nous aurons besoin aujourd'hui.