# Semaine 4

## Consignes

Avant le lundi 21 février 22h00:
- [x] Relire les [diapos du cours 4](CM4.md)
- [x] Rendu d'étape du TP (voir ci-dessous)

## Cours

- [X] [CM4: Construction et sélection d'attributs](CM4.md)

## TP: VI-ME-RÉ-BAR sur vos propres données

### Objectif

La semaine dernière, vous vous êtes familiarisé avec le schéma
VI-ME-RÉ-BAR -- [VI]sualisation, [ME]sure, [RÉ]férence (*baseline*),
[BAR]res d'erreur -- en reproduisant une analyse de données
préexistante dont l'objet était de classifier automatiquement des
images de pommes et bananes. Maintenant, c'est à vous de jouer! Vous
allez effectuer en binôme votre propre analyse de données.

La semaine dernière, vous avez choisi l'un des [jeux de données
fournis](../Semaine3/3_jeux_de_donnees.md) consistant chacun
de vingt images réparties en deux classes: poules et canards,
mélanomes cancéreux et bénins, émoticons tristes et gais, paumes et
dos de la main, ou chiffres manuscrits zéro et un.

Saurez-vous apprendre à l'ordinateur à distinguer automatiquement, par
exemple, un mélanome cancéreux d'un mélanome bénin, uniquement à
partir de sa photographie? Nous avons veillé à ce que votre défi ne
soit ni trop simple, ni trop compliqué.

Ce travail sera l'objet du premier projet qui va se
dérouler sur les deux semaines qui viennent:

Semaine 4: dépôt lundi 21 février 22h
- Choix définitif du binôme (voir annonce sur e-Campus et [document partagé](https://codimd.math.cnrs.fr/3f98v4-YT4CN2ktM6_g5lw?both))
- Traitement des images numériques pour en extraire des attributs
  (*features*) simples.
- Utilisation d'un premier classificateur donné par vos enseignantes et
  enseignants, afin d'obtenir une performance de [RÉ]férence
  (*baseline*) pour ce jeu de données.
- Évaluation de la performance.

Semaine 5: dépôt mardi 1er mars 22h
- Implémentation d'un classificateur.
- Évaluation de la performance.

Vous déposerez chaque semaine votre travail dans son état
d'avancement. Votre version définitive sera évaluée par vos
enseignants (20% de la note finale).

### Consignes

Vous documenterez au fur et à mesure votre analyse de données dans le
document exécutable [analyse de données](4_analyse_de_donnees.md), en
suivant la trame fournie. Gardez notamment une trace des
expérimentations intermédiaires («nous avons essayé telle combinaison
d'attributs; voici les résultats et la performance obtenue»). Ce
document devra rester à tout moment synthétique, suivant notamment les
[bonnes pratiques](../Semaine3/1_bonnes_pratiques.md) vues la semaine
dernière:
- Vous mettrez dans le fichier `utilities.py` les utilitaires du TP3
  (`load_images`, ...)  que vous souhaiterez réutiliser, ainsi que vos
  nouvelles fonctions.
- Complétez régulièrement le rapport ci-dessous, notamment pour qu'il
  affiche le code de toutes les fonctions que vous avez
  implantées. Vérifiez à chaque fois le résultat des outils de
  vérifications (`flake8`, ...).
- Lorsque vous aurez besoin de brouillon -- par exemple pour mettre au
  point du code -- créez des petites feuilles Jupyter séparées pour ne
  pas polluer votre document.

La qualité de la rédaction sera l'un des critères d'évaluation du
mini-projet.

### Au travail!

- [X] Vérifiez votre inscription avec votre binôme pour le projet 1
     dans le [document
     partagé](https://codimd.math.cnrs.fr/3f98v4-YT4CN2ktM6_g5lw?both).
     Inscrivez-vous aussi si vous n'avez pas encore de binôme!
- [X] Téléchargez le sujet de TP `Semaine4` (rappel des
     [instructions](http://nicolas.thiery.name/Enseignement/IntroScienceDonnees/Devoirs/Semaine1/index.html#telechargement-et-depot-des-tps))
- [X] Ouvrez la feuille [index](index.md) pour retrouver ces consignes.
- [X] Consultez la section « Rapport » en fin de feuille.
- [X] Faites un rappel sur la [manipulation des tableaux](1_tableaux.md).
- [X] Apprenez à traiter des [images](2_images.md)  ...
- [X] et à [extraire des attributs](3_extraction_d_attributs.md) de
     votre jeu de données
- [X] Effectuez votre [analyse de donnees](4_analyse_de_donnees.md),
     en suivant les instructions fournies.

### Brouillon de rapport

Cette feuille joue aussi le rôle de brouillon de mini-rapport qui vous
permettra à vous et votre enseignant d'évaluer rapidement votre
avancement sur cette première partie du projet. Cela vous donnera une
base pour le mini-rapport de projet la semaine prochaine.

Au fur et à mesure du TP, vous cocherez ci-dessus les actions que vous
aurez effectuées; pour cela, double-cliquez sur la cellule pour
l'éditer, et remplacez `- [ ]` par `- [x]`. Vous prendrez aussi des
notes ci-dessous. Enfin, vous consulterez la section « Revue de code »
ci-dessous pour vérifier la qualité de votre code.

<!-- grade_id: cell-4911a792a82448dF, points: 0, grade: true, solution: true -->
Lien vers le rapport: https://github.com/MK8BK/SDprojet_image_s2

#### Revue du code

##### Affichage du code des principales fonctions

```python
from utilities import *
# Feuille 2_images.md
show_source(show_color_channels)
```

##### Conventions de codage

L'outil `flake8` permet de vérifier que votre code respecte les
conventions de codage usuelles de Python, telles que définies
notamment par le document [PEP
8](https://www.python.org/dev/peps/pep-0008/). Si la cellule suivante
affiche des avertissements, suivez les indications données pour
peaufiner votre code.

<!-- grade_id: cell-ce5cf24a9e83ff0e, points: 2, grade: true, solution: false -->
```python
assert run_without_error("flake8 utilities.py")
```