# Semaine 5

## Consignes

Avant mardi 1er mars à 22h:
- Relire les [diapos du cours 5](CM5.md)
- Rendu final du mini-projet 1. Vous devrez avoir déposé votre version
  finale du dossier `Semaine5` sur GitLab. Pour les binômes, chaque
  membre devra faire un dépôt séparé, même si les deux soumissions
  sont identiques.

## TP: fin du mini-projet 1

La semaine précédente, vous avez appliqué le schéma VI-ME-RÉ-BAR --
[VI]sualisation, [MÉ]trique, [Ré]férence, [BAR]res d'erreur -- pour la
classification de votre propre jeu de données. Vous avez:

- Fait quelques essais avec un des jeux de données fournis
- Déroulé l'analyse de données avec des attributs simples (par exemple
  rougeur et élongation) ainsi qu'un premier classifieur, afin
  d'obtenir votre référence: sans chercher à optimiser, quelle
  performance de classification obtient-on?
- Implémenté d'autres attributs et déroulé à nouveau l'analyse pour
  obtenir une meilleure performance.

### Objectifs

L'objectif de cette deuxième semaine de mini-projet est de mettre en
application le [cours de cette semaine](CM5.md) en explorant d'autres
classificateurs. Vous pouvez au choix:

- Implémenter l'un des classificateurs décrits dans le cours:
  - OneR(**)
  - KNN: k-plus proches voisins (*)
  - Arbre de décision (*)
  - Fenêtre de Parzen (*)
  - Perceptron (**)
- Essayer plusieurs autres [classificateurs fournis par Scikit-Learn](https://scikit-learn.org/stable/auto_examples/classification/plot_classifier_comparison.html) comme par exemple:
  - Autre méthode de noyau (****)
  - Perceptron multi-couche (****)

### Au travail!

- [ ] Vérifiez votre inscription avec votre binôme pour le projet 1
     dans le [document
     partagé](https://codimd.math.cnrs.fr/3f98v4-YT4CN2ktM6_g5lw?both).
     Inscrivez-vous aussi si vous n'avez pas encore de binôme!
- [ ] Téléchargez le sujet de TP `Semaine5` (rappel des
     [instructions](http://nicolas.thiery.name/Enseignement/IntroScienceDonnees/Devoirs/Semaine1/index.html#telechargement-et-depot-des-tps))
- [ ] Recopiez dans `Semaine5` vos feuilles de travail de la semaine 4:

        cd ~/IntroScienceDonnees/Semaine5
        cp ../Semaine4/?_*.md .

- [ ] Ouvrez la feuille [index](index.md) pour retrouver ces consignes.
- [ ] Consultez la section « Rapport » en fin de feuille.
- [ ] Révisez les [bonnes pratiques](../Semaine3/1_bonnes_pratiques.md)
	 vues en semaine 3.
- [ ] Partez à la découverte des [classificateurs](5_classificateurs.md).
- [ ] Reprenez votre analyse de la semaine dernière dans la feuille
     [analyse de donnees](4_analyse_de_donnees.md) avec vos propres
     classificateurs.


- [ ] Si nécessaire, relisez les
     [instructions](http://nicolas.thiery.name/Enseignement/IntroScienceDonnees/Students/Semaine1/index.html#telechargement-et-depot-des-tps)
     pour le téléchargement et le dépôt des TPs, ainsi que les 
	 [bonnes pratiques](http://nicolas.thiery.name/Enseignement/IntroScienceDonnees/Students/Semaine1/index.html#bonnes-pratiques).

### Rapport

Cette feuille joue le rôle de mini-rapport ci-dessous. Elle vous
permettra à vous et votre enseignant ou enseignante d'évaluer rapidement votre
avancement sur cette première partie du projet. 

Au fur et à mesure, vous cocherez ci-dessus les actions que vous aurez
effectuées; pour cela, double-cliquez sur la cellule pour l'éditer, et
remplacez `- [ ]` par `- [x]`. Vous prendrez aussi des notes
ci-dessous. Enfin, vous consulterez la section « Revue de code »
ci-dessous pour vérifier la qualité de votre code.

<!-- grade_id: cell-4911a792a82448dF, points: 1, grade: true, solution: true -->
- [x] Vérifiez votre inscription avec votre binôme pour le projet 1
     dans le [document
     partagé](https://codimd.math.cnrs.fr/3f98v4-YT4CN2ktM6_g5lw?both).
     Inscrivez-vous aussi si vous n'avez pas encore de binôme!
- [x] Téléchargez le sujet de TP `Semaine5` (rappel des
     [instructions](http://nicolas.thiery.name/Enseignement/IntroScienceDonnees/Devoirs/Semaine1/index.html#telechargement-et-depot-des-tps))
- [x] Recopiez dans `Semaine5` vos feuilles de travail de la semaine 4:

        cd ~/IntroScienceDonnees/Semaine5
        cp ../Semaine4/?_*.md .

- [x] Ouvrez la feuille [index](index.md) pour retrouver ces consignes.
- [x] Consultez la section « Rapport » en fin de feuille.
- [x] Révisez les [bonnes pratiques](../Semaine3/1_bonnes_pratiques.md)
	 vues en semaine 3.
- [ ] Partez à la découverte des [classificateurs](5_classificateurs.md).
- [ ] Reprenez votre analyse de la semaine dernière dans la feuille
     [analyse de donnees](4_analyse_de_donnees.md) avec vos propres
     classificateurs.


- [ ] Si nécessaire, relisez les
     [instructions](http://nicolas.thiery.name/Enseignement/IntroScienceDonnees/Students/Semaine1/index.html#telechargement-et-depot-des-tps)
     pour le téléchargement et le dépôt des TPs, ainsi que les 
	 [bonnes pratiques](http://nicolas.thiery.name/Enseignement/IntroScienceDonnees/Students/Semaine1/index.html#bonnes-pratiques).

#### Revue du code

##### Affichage du code des principales fonctions

```python
from utilities import *
# Feuille 2_images.md
show_source(show_color_channels)
```

Output:
```
<IPython.core.display.HTML object>
```

##### Conventions de codage

L'outil `flake8` permet de vérifier que votre code respecte les
conventions de codage usuelles de Python, telles que définies
notamment par le document [PEP
8](https://www.python.org/dev/peps/pep-0008/). Si la cellule suivante
affiche des avertissements, suivez les indications données pour
peaufiner votre code.

<!-- grade_id: cell-ce5cf24a9e83ff0e, points: 2, grade: true, solution: false -->
```python
assert run_without_error("flake8 utilities.py")
```

Output:
```
NameError                                 Traceback (most recent call last)
~\AppData\Local\Temp/ipykernel_15412/3929026667.py in <module>
----> 1 assert run_without_error("flake8 utilities.py")

NameError: name 'run_without_error' is not defined
```

### Barême indicatif /20

* 1_tableaux.md : 1,5 points 
* 2_images.md : 2 points 
* 3_features.md : 1 point
* 4_analyse_de_données.md : 10 points 
* 5_classificateur.md : 4 points (+ 3 points bonus pour le OneR)
* index (semaine 4 et 5) : 1,5 points

```python

```